Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Linh Truong"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 11.1

The pre-class work for this session will focus on the rod cutting problem. Recall that the rod cutting problem takes as an input the length n of a rodand a table of prices $p_i$ for $i = 1,2,... n$, and one needs to determine the maximum revenue $r$ obtainable by cutting up the rod and selling the pieces. 

## Part A. 

You will implement in Python two solutions to the rod cutting problem, namely:

## Question 1. 
A recursive top-down solution to the rod cutting problem. Please complete the cut_rod function below:


<!-- grade_id: cell-d85ce9860c9ea1e5, points: 0, grade: false, solution: true -->
```python
def cut_rod(p,n):
    """
    A recursive top-down solution to the rod cutting problem as described in 
    Cormen et. al. (pp 363) This calculates the maximum revenue that could be 
    earned by cutting up a rod of length n.
    
    Inputs;:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the dollars
    of revenue the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - q: float, the optimal revenue
    """
    if n == 0:
        return 0
    q = -1000
    for i in range(1,n+1): #taking a rod of length i
        q = max(q,p[i]+ cut_rod(p,n-i)) 
    return q

  
```

<!-- grade_id: cell-0288a4047be5088d, points: 1, grade: true, solution: false -->
```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1

#results from textbook
r = [0,1,5,8,10,13,17,18,22,25,30]
for i in range(len(r)):
    assert(cut_rod(p,i)==r[i])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-3-688fa9e79295> in <module>
      5 r = [0,1,5,8,10,13,17,18,22,25,30]
      6 for i in range(len(r)):
----> 7     assert(cut_rod(p,i)==r[i])

AssertionError:
```

## Question 2.

An optimal rod cutting solution using dynamic programming (see figure [here](https://drive.google.com/open?id=1nu9gETKX4KJCHZDi17fZLQtkDVig81Zk) for inspiration). Please complete the following two functions: 


<!-- grade_id: cell-eb49e8516ebf80aa, points: 0, grade: false, solution: true -->
```python
def extended_bottom_up_cut_rod(p,n):
    """
    Implements a bottom-up dynamic programming approach to the rod cutting problem.
    Here, "extended" means the function is geared in a way amenable to reconstructing
    an optimal solution, on top of the returned optimal value. See Cormen et al.,
    p. 269 for the implementation details.
    
    Inputs:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the dollars
    of revenue the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - r: list of floats, the maximum revenues. r[i] gives the maximum revenue for a rod
    of length i. As such:
        * r[0] = 0
        * len(r) == n + 1
    - s: list of ints, the optimal sizes of the first piece to cut off. Also make sure 
    that:
        * s[0] = 0
        * len(s) == n + 1
    """
    r = [i for i in range(n+1)] #storing best result
    s = [i for i in range(n+1)] #storing solutions
    r[0] = 0
    
    for j in range(1,n+1):
        q = -1000
        for i in range(1,j+1):
            if q < p[i] + r[j-i]: 
                q = p[i] + r[j-i]
                s[j] = i #optimal size of 1st piece to cut off
        r[j] = q #optimal revenue    
    return r, s
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1d52c9519ccb3502, points: 0, grade: false, solution: true -->
```python
def print_cut_rod_solution(p,n):
    """
    Gives a solution to the rod cutting problem of size n. 
    
    Inputs:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the revenue (in USD, for example) the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - sol: a list of ints, indicating how to cut the rod. Cutting the rod with the lengths
    given in sol gives the optimal revenue.
        * print_cut_rod_solution(p,0) == []
    """
    (r,s) = extended_bottom_up_cut_rod(p,n)
    while n > 0:
        print(s[n])
        n = n - s[n]
    return 

```

<!-- grade_id: cell-6575aa2769e0eab1, points: 1, grade: true, solution: false -->
```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1
# Result of R and S from textbook:
R = [0,1,5,8,10,13,17,18,22,25,30]
S = [0,1,2,3,2,2,6,1,2,3,10]
# Test:
r, s = extended_bottom_up_cut_rod(p,10)
assert(r==R)
assert(s==S)
```

Output:
```
<ipython-input-6-0a1b9fc5af6b> in <module>
      5 S = [0,1,2,3,2,2,6,1,2,3,10]
      6 # Test:
----> 7 r, s = extended_bottom_up_cut_rod(p,10)
      8 assert(r==R)
      9 assert(s==S)

<ipython-input-4-2a1319d79540> in extended_bottom_up_cut_rod(p, n)
     28         q = -1000
     29         for i in range(1,j+1):
---> 30             if q < p[i] + r[j-i]:
     31                 q = p[i] + r[j-i]
     32                 s[j] = i #optimal size of 1st piece to cut off

IndexError: list index out of range
```

## Part B - Experiments

## Question 1.

Use the function below to generate a list of prices of length n=20 and assign that list to a new variable, `P`. You MUST use this list for parts 2 and 3 below.


```python
import numpy as np
def generate_price_list(n):
    """Generates a price list of length n
    
    Inputs:
    - n: integer, length of the list, must be positive
    
    Outputs:
    - p: list, the ordered price list for each rod cut
    """
    p = [1]
    for i in range(1,n):
        np.random.seed(0)
        p.append(np.random.randint(p[i-1]+1, i*3+1))
    return p

```

## Question 2.

Time (using the time library) the implementation of `cut_rod` for different rod lengths for values of up to $n=20.$ Present your results in a plot. 


<!-- grade_id: cell-e094e913d8bc139f, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 3.

Time (using the time library) the implementation of `extended_bottom_up_cut_rod` for different rod lengths for values of up to $n=20$. Add the curve corresponding to this algorithm to the previous plot.

<!-- grade_id: cell-210aa3d5a2743bd8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 4.

Contrast both curves. What can you conclude?

<!-- grade_id: cell-0e62cc5b19840488, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Part C - After completing your experiments above, answer the following questions:

## Question 1.
Estimate how big $n$ has to be before the non-dynamic-programming version will take a million years to finish. Call this value N.

<!-- grade_id: cell-a11f4fd6a03e2cf5, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.
Estimate (or time) how long it takes to evaluate `extended_bottom_up_cut_rod(p,N)`, where `N` is the value you got from the previous question. 

<!-- grade_id: cell-89ed2dd6f3310c96, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## [Optional]  Question 3. 
Do you notice anything peculiar about the solution for large n?

<!-- grade_id: cell-fd4c9c91a42f178a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE