Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Linh Phuong Truong"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    maximum = 0
    for i in range(1, len(A)):
        for j in range(i, len(A)):
            current_sum = 0
            for k in range(i, j):
                current_sum += A[k]
            maximum = max(maximum, current_sum)
    return maximum
A = [-2,-3,4,-1,-2,1]
bruteforce_max_subarray(A)

```

Output:
```
4
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
With the previous code, I will have to run the entire algorithm again if a number is added into the list. This is because the previous algorithm disregards the previous effort made to find the maximum subarrays, thus if there is a new element, the whole process needs to be carried out again to find the maximum subarray.

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False. If the max subarray is different when a new element is introduced, the chances are that the new max subarray has the newly apprehended element.

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    maximum = 0
    for i in range(1, len(A)):
        for j in range(i, len(A)):
            current_sum = 0
            for k in range(i, j):
                current_sum += A[k]
            maximum = max(maximum, current_sum)
    return maximum
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    # YOUR CODE HERE
    if len(x)==1:
        return x
    else:
        
        #x is an array of integers
        #mx is the value of the maximum subarray of x[0:-1] which then returns the max_subarray of the entire list
        #function output is the maximum sum of the subarray
        max_sum = 0
        s = 0
        if x[-1] <= 0:
            max_sum = mx #if the last item is smaller or equal to 0, the maximum subarray will remain the same.
        else:
            i = 0
            for i in range(len(x)):
                m = len(x) - i -1
                s = s + x[m]
                if s > mx:
                    max_sum = s
    return max_sum
    raise NotImplementedError()
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-5-2256e7bbc067> in <module>
      1 B = [-2, -3, 4, -1, -2, 1, 8]
----> 2 assert(incremental_max_subarray(B, 4) == 10)
      3 assert(incremental_max_subarray(B[:1], 0) == B[0])

AssertionError:
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE