Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Linh Phuong Truong"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random
def sort_check(A):
    n = len(A)
    for i in range (0,n-1):
        if A[i] > A[i+1]:
            return False
    return True

def shuffle(A):
    n = len(A)
    for i in range (0,n):
        j = random.randint (0, n-1)
        A[i], A[j] = A[j], A[i]

def worst_sort(A):
    n = len(A)
    while (sort_check(A) == False):
        shuffle(A)
    return A    

a = [5555,2,-89,36,47,2,0] 
worst_sort(a)
```

Output:
```
[-89, 0, 2, 2, 36, 47, 5555]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The best case complexity is the case when on the first shuffle, the algorithm generate the correctly ascending order list. It leads to O(n) - time taken to check all n variables.

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
    - There are n! permutation for an input of n.
    - Since we are randomizing the elements in the list, they have inform distribution. The probability of having the correct shuffle is 1/n!. 
    - The average number of attempt is O(n!)
    - It takes O(n) time each shuffle to check if it is the right list. Thus the average case complexity is O(n*n!)

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
The list should be as small in size as possible to reduce running time. Hypothetically a big list is still possible, it just takes forever to run based on the complexity analysis above.

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.
Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-δ/2)% and (50+δ/2)% percentiles. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
import numpy as np
import random
def check_approx_med(A, num, delta):
    n = len(A)
    low = np.percentile(A,(50-delta/2))
    high =  np.percentile(A,(50+delta/2))
    for i in range (10):
        result = check(num)
        if result == True:
            return result
        else:
            new(A)
    return result

def check(num):
    if (num >high) or (num<low):
        return False
    else:
        return True

def new(A):
    n = len(A)
    num = random.randint (0, n-1)
    return num
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
import numpy as np
import random
def approx_med_finder(A, delta):
    n = len(A)
    for i in range (100//delta):
        num = new(A)
        result = check_approx_med(A,num,delta)
        if result == True:
            return num
    return result

def new(A):
    n = len(A)
    num = random.randint (0, n-1)
    return num

    raise NotImplementedError()
```

```python
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    # YOUR CODE HERE
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
Since there are n element in the list, and the probability of failure is 1-((50+delta/2)-(50-delta/2))/n = (1-delta)/n.
After all allowed random trials:
100/delta * (1-delta)/n 

I think it scale with delta

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
check_approx_med retries only 10 times, each of which it will have to generate a number and check to see if it is in the range, so I think it is O(10*2)
approx_med_finder tries it 100/delta times, thus its complexity is 10000/O(delta)