Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Linh Truong"
COLLABORATORS = "Elena Cai"
```

---

# CS110 Pre-class Work 7.1

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson 3.2, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1.
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


```python
import heapq

def max_heapify(A):
    A_negative = []
    for i in A:
        A_negative.append(-i)
    A = []
    heapq.heapify(A_negative)
    for i in A_negative:
        A.append(-i)
    return(A)

def min_heapify(A):
    heapq.heapify(A)
    return(A)
```

<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    
    #determine the current median
    if len(minh) > len(maxh):
        median = minh[0]
    elif len(minh) < len(maxh):
        median = maxh[0]
    elif len(minh) == len(maxh) and len(minh)!=0:
        median = (minh[0]+maxh[0])/2
    elif len(minh) == len(maxh) == 0:
        median = 0
    
    #If the new element is greater than the current median, it goes to the min-heap. 
    if elem > median:
        minh.append(elem)
        minh = min_heapify(minh)
        
    #If it is less than the current median, it goes to the max heap.
    if elem <= median:
        maxh.append(elem)
        maxh = max_heapify(maxh)
        
    #If imbalanced, put one element to the other list
    if len(minh) > len(maxh)+1:
        min_first = []
        min_first.append(minh[0])
        del minh[0]
        maxh = min_first + maxh
        
    if len(maxh) > len(minh)+1:
        max_first = []
        max_first.append(maxh[0])
        del maxh[0]
        minh = max_first + minh
    
    return minh, maxh
```

```python
add_to_median_heap([7,8,9],[5,4,2,3],3)
```

Output:
```
([5, 7, 8, 9], [4, 2, 3, 3])
```

## Question 2
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
#return the median element of two heaps, with a maximum heap size difference of 1

def median(minh, maxh):
    if len(minh) > len(maxh):
        median = minh[0]
    elif len(minh) < len(maxh):
        median = maxh[0]
    elif len(minh) == len(maxh) and len(minh)!=0:
        median = (minh[0]+maxh[0])/2
    elif len(minh) == len(maxh) == 0:
        median = 0
    return median
```

```python
median([7,8,9],[5,4,2,3])
```

Output:
```
5
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3.

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
# minh = []
# maxh = []
# for a in range(1,100,2):
#     add_to_median_heap(minh, maxh, a)
#     print(median(minh, maxh))  
```

```python
minh = []
maxh = []
for a in range(1,100,2):
    minh,maxh = add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
2.0
3
4.0
5
6.0
7
8.0
9
10.0
11
12.0
13
14.0
15
16.0
17
18.0
19
20.0
21
22.0
23
24.0
25
26.0
27
28.0
29
30.0
31
32.0
33
34.0
35
36.0
37
38.0
39
40.0
41
42.0
43
44.0
45
46.0
47
48.0
49
50.0
```

## Question 4.
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
Heapify has O(lg n). Worst case is when the new element is put into the heap, the other heap is still empty, thus the algorithm needs to heapify a list and move on element. Overall, a list of n would take O(nlogn)

## Question 5.
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
I think it is O(n), when we need to heapify each small list. But if the inputed list is sorted then O(3)

## Question 6.

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
Depends on the sorting list algorithm we use, but the worst case of this algorithm is n log n which is not too bad.

## [Optional] Question 7.

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1.

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


```python
import random

def qselect(lst,k):
    
    partition = random.choice(lst)
    
    lower = [a for a in lst if a < partition] #elements smaller than the partition
    len_l = len(lower)

    middle = [a for a in lst if a == partition] #elements equal to the partition
    len_m = len(middle)
    
    upper = [a for a in lst if a > partition] #elements greater than the partition
    len_u = len(upper)
    
    #if k-th smallest is in lower
    if k < len_l:
        return qselect(lower,k)
    
    #if k-th smallest is in upper
    elif k > len_l+len_m:
        return qselect(upper,k-len_l-len_m)
        
    #if k-th smallest is the partition
    else:
        return(partition)
```

```python
qselect([1,2,3,20,20,30,10,30,20],8)
```

Output:
```
30
```

<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Output:
```
File "<ipython-input-13-dd00d98f476b>", line 2
    # YOUR CODE HERE
                    ^
SyntaxError: unexpected EOF while parsing
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
<ipython-input-14-253afcee666d> in <module>
      5 lst2 = []
      6 for a in range(100):
----> 7     lst2.append(qselect(lst1, a))
      8 assert(lst2 == list(range(100)))

<ipython-input-11-1190b9902ffb> in qselect(lst, k)
      1 def qselect(lst,k):
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

## Question 2.
Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.


```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
     print(qselect(lst, a))

```

Output:
```
0
1
2
2
3
4
5
6
7
8
9
11
12
13
13
14
16
16
18
19
19
20
21
23
23
25
26
26
27
29
29
31
31
33
33
35
35
36
37
39
39
40
42
43
44
45
46
46
47
48
50
51
52
52
54
54
55
56
58
59
59
60
61
63
64
65
65
66
68
68
70
71
71
72
74
74
75
77
78
79
80
81
82
82
84
84
85
87
87
89
89
90
91
93
94
94
96
96
98
99
```

## Question 3.

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
T(n) = T(n/2) + c

## Question 4.

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
Best case O(1). This is when the first step returns the index. 


## Question 5.

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
O(n^2) When each step only delete one element of the list. It would take n steps in total. With each step n.  
