Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Linh Phuong Truong"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.2

## Question 1.
Using Figure 7.1 in Cormen et al. as a model, perform manually the partition process on the following list: A = [1,5,6,2,3,8,9,4,7]. You just need to specify the followings:
1. The array after the process is done.
2. The value of $i$ after the process is done.

<!-- grade_id: cell-7aa520f8af13679b, points: 0, grade: true, solution: true -->
1. [1,5,6,2,3,4,7,8,9]
2. i = 5 (starting with i = -1)

## Question 2.
Code up a Python implementation of `partition(A, p, r)`, closely follow the pseudo-code in Cormen et al., p.172. Your function should return the index of the pivot in the array.

<!-- grade_id: cell-1ceb2600756c60ff, points: 0, grade: false, solution: true -->
```python
def partition(A,p,r):
    i = p - 1
    x = A[r]
    for j in range (p,r-1):
        if A[j] <= x:
            i += 1
            A[i],A[j] = A[j],A[i]
        A[i+1],A[r] = A[r],A[i+1]
    return (i+1)
    """
    Assume r<len(A) and p>=0
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a57b60117a7b82fb, points: 1, grade: true, solution: false -->
```python
A = [1,5,6,2,3,8,9,4,7]
assert(partition(A, 0, len(A)-1)==6)
```

## Question 3.

Code up your own Python implementation of `quicksort(A, p, r)`, using `partition(A,p,r)`.

<!-- grade_id: cell-8c39ebb8cd1aa83a, points: 0, grade: false, solution: true -->
```python
def quick_sort(A,p,r):
    if p < r:
        q = partition(A,p,r)
        quick_sort(A,p,q-1)
        quick_sort(A,q+1,r)
    else:
        return A
    raise NotImplementedError()

```

Output:
```
File "<ipython-input-52-f1caae2f1403>", line 11
SyntaxError: 'break' outside loop
```

<!-- grade_id: cell-4f822430efd456ee, points: 1, grade: true, solution: false -->
```python
A = [0]
assert(quick_sort(A, 0, 0) == [0])
A = [3,1,2]
assert(quick_sort(A, 0, 2) == [1,2,3])
```

Output:
```
<ipython-input-50-34c0f35e8a7b> in quick_sort(A, p, r)
      3         q = partition(A,p,r)
      4         quick_sort(A,p,q-1)
----> 5         quick_sort(A,q+1,r)
      6     else:
      7         return A

<ipython-input-50-34c0f35e8a7b> in quick_sort(A, p, r)
      6     else:
      7         return A
----> 8     raise NotImplementedError()

NotImplementedError:
```

## Question 4. 
Explain (using experimental plots) the running time of `quick_sort` when: 
1. all elements of array A have the same value (e.g., [1,1,1])?
2. array A contains distinct elements sorted in decreasing order (e.g., [5,4,2,1])?


<!-- grade_id: cell-b58035dd5fa02329, points: 0, grade: true, solution: true -->
```python
A=[1,1,1]
quick_sort(A)
plot(A)
raise NotImplementedError()
```